```python
import numpy as np
from helpers import print_table
```

# 2. Pairwise vs. Mutual Independence (6 points)

**Definition**: We say that two random variables are *pairwise independent* if $$p(X_n \mid X_m) = p(X_n)$$ and hence $$p(X_m, X_n) =  p(X_n \mid X_m) p(X_m) = p(X_n) p(X_m) $$

**Definition**: We say that $n$ random variables are *mutually independent* if $$p(X_i \mid X_{S}) = p(X_i)\;\; \forall S \subseteq \{1, \dots, n\} \setminus \{ i \}$$ and hence $$p(X_{1:n}) = \prod_{i=1}^n p(X_i)$$




<div class="alert alert-warning">
Show that pairwise independence between all pairs of variables does not necessarily imply mutual independence. Come up with a minimal counter example that has exactly three binary random variables.
</div>

Specify this counterexample via its full joint distribution table (FJDT). **Briefly** outline your thought process in the text field below (use $\LaTeX$ and markdown) and store the model's full joint distribution table into the `XYZ` variable. It is sufficient to show pairwise independence and non-mutual independence by comparing products of marginals and joint distributions. 

**Hint**: Copy your implementation of `inference_by_enumeration` from Problem 1. You can use `print_table` to visualize your distribution tables such as the FJDT, products of marginals, and joint distributions.

```python
help(print_table)
```

Output:
```
Help on function print_table in module helpers:

print_table(probability_table: numpy.ndarray, variable_names: str) -> None
    Prints a probability distribution table.

    Parameters
    ----------
    probability_table : np.ndarray
        The probability distribution table
    variable_names : str
        A string containing the variable names, e.g., 'CDE'.

    Returns
    -------
    None
```

<!-- grade_id: cell-6ced5d0ec1cbd764, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

as an example, we can assume we have a probabilistic case with these conditions:
P(X = 0, Y = 0, Z = 0) = 1/4
P(X = 1, Y = 1, Z = 0) = 1/4
P(X = 1, Y = 0, Z = 1) = 1/4
P(X = 0, Y = 1, Z = 1) = 1/4
P(X = 0, Y = 1, Z = 0) = 0
P(X = 0, Y = 1, Z = 1) = 0
P(X = 1, Y = 1, Z = 1) = 0
P(X = 1, Y = 0, Z =
 so from these conditions, we can conclude that:0) = 0

P(X) = P(Y) = P(Z) = 1/2
P(X Ո Y) = P(X Ո Z) = P(Y4
 Z) = 1/2
P(X Ո Y so:
P(X) * P(Y) = 1/4
P(X) * P(Z) = 1/4
P(Y) * P(Z) = 1/4
Ո Z) = 1/4
P(X) * P(Y) 

so we can see that [P(X) * P(Y) = P(X Ո Y), P(X) * P(Z) = P(X Ո Z), P(Y) * P(Z) = P(Y Ո Z)] which is an indication of pairwise independence,
but [P(X Ո Y Ո Z)] does not equal [P(X) * P(Y) * P(Z)] so it is not mutually independent.* P(Z) = 1/8


<!-- grade_id: cell-7f0b4fa0aa65bb39, points: 0, grade: false, solution: true -->
```python
XYZ = np.ones((2,2,2))

# YOUR CODE HERE


XYZ[0,0,0] = 1/4
XYZ[0,1,1] = 1/4
XYZ[1,0,1] = 1/4
XYZ[1,1,0] = 1/4
XYZ[0,1,0] = 0
XYZ[0,0,1] = 0
XYZ[1,0,0] = 0
XYZ[1,1,1] = 0

print("FJDT")
print_table(XYZ, "XYZ")
#raise NotImplementedError()
```

Output:
```
FJDT
<IPython.core.display.HTML object>
```

```python
def inference_by_enumeration(
    FJDT: np.ndarray, 
    query_variable_indices: tuple, 
    evidence_variable_indices: tuple = tuple()
) -> np.ndarray:
 
    assert type(FJDT) == np.ndarray, "FJDT must be a np.ndarray"
    assert type(query_variable_indices) == tuple, "query_variable_indices must be a tuple"
    assert type(evidence_variable_indices) == tuple, "evidence_variable_indices must be a tuple"
        
    query_sum = query_variable_indices + evidence_variable_indices
    hv = tuple(set(range(FJDT.ndim)).difference(query_sum))
    p = np.sum(FJDT, axis=hv, keepdims=True)
    z = np.sum(p, axis=query_variable_indices, keepdims=True)
    CPT = 1/z * p
    #print(CPT)

    return CPT
```

<!-- grade_id: cell-1486738b77c0d868, points: 0, grade: false, solution: true -->
```python
# copy inference_by_enumeration from Problem 1 & print and compare the probability tables here!

# YOUR CODE HERE
X, Y, Z = 0, 1, 2

P_X = inference_by_enumeration(XYZ, (X,), ())
P_Y = inference_by_enumeration(XYZ, (Y,), ())
P_Z = inference_by_enumeration(XYZ, (Z,), ())
P_XY = inference_by_enumeration(XYZ, (X,Y), ())
P_XZ = inference_by_enumeration(XYZ, (X,Z), ())
P_YZ = inference_by_enumeration(XYZ, (Y,Z), ())
P_XYZ_jont = inference_by_enumeration(XYZ, (X,Y,Z), ())
p_XYZ_product = P_X * P_Y * P_Z

print('P(X):')
print_table(P_X.squeeze(),"X")
print('P(Y):')
print_table(P_Y.squeeze(),"Y")
print('P(Z):')
print_table(P_Z.squeeze(),"Z")
print('P(X Ո Y):')
print_table(P_XY.squeeze(),"XY")
print('P(X)*P(Y):')
print_table((P_X*P_Y).squeeze(),"XY")
print('P(X Ո Z):')
print_table(P_XZ.squeeze(),"XZ")
print('P(X)*P(Z):')
print_table((P_X*P_Z).squeeze(),"XZ")
print('P(Y Ո Z):')
print_table(P_YZ.squeeze(),"YZ")
print('P(Y)*P(Z):')
print_table((P_Y*P_Z).squeeze(),"YZ")
print('P(X Ո Y Ո Z):')
print_table(P_XYZ_jont.squeeze(),"XYZ")
print('P(X)*P(Y)*P(Z):')
print_table((p_XYZ_product).squeeze(),"XYZ")
#raise NotImplementedError()
```

Output:
```
P(X):
<IPython.core.display.HTML object>
P(Y):
<IPython.core.display.HTML object>
P(Z):
<IPython.core.display.HTML object>
P(X Ո Y):
<IPython.core.display.HTML object>
P(X)*P(Y):
<IPython.core.display.HTML object>
P(X Ո Z):
<IPython.core.display.HTML object>
P(X)*P(Z):
<IPython.core.display.HTML object>
P(Y Ո Z):
<IPython.core.display.HTML object>
P(Y)*P(Z):
<IPython.core.display.HTML object>
P(X Ո Y Ո Z):
<IPython.core.display.HTML object>
P(X)*P(Y)*P(Z):
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b612791dd41278bf, points: 4, grade: true, solution: false -->
```python
assert XYZ.shape == (2, 2, 2), 'FJDT must have shape (2,2,2)'
assert np.isclose(XYZ.sum(), 1), 'Probabilites in FJDT must sum to one'

```

```python

```

```python

```