---
# 3. Conditional Independence (4 points)

Let $Y$ be a discrete random variable over $K$ classes $\{1, \dots, K\}$ and $X_1, X_2 \in \mathbb{R}$ two continuous random variables (features). Suppose we wish to calculate the posterior distribution over classes given observation $(x_1, x_2)$, i.e., $$P(Y \mid x_1, x_2) = \big[P(Y = 1 \mid x_1, x_2), \dots, P(Y = K \mid x_1, x_2) \big]$$

---
## 3.1. Posterior Over Classes

Which of the following numbers / probability distribution tables are sufficient to compute the posterior distribution $P(Y \mid x_1, x_2)$? Select all answers that apply!
1. $P(x_1, x_2)$, $P(Y)$, $P(x_1 \mid Y), P(x_2 \mid Y)$
2. $P(x_1, x_2)$, $P(Y)$, $P(x_1, x_2 \mid Y)$
3. $P(x_1 \mid Y)$, $P(x_2, \mid Y)$, $P(Y)$

<div class="alert alert-warning">
Show your calculations in the text field (use markdown & $\LaTeX$) and indicate your solutions in the code cell below. **Keep your solution short**! (2 points)
</div>

<!-- grade_id: cell-88f39e958328a739, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

since we need to compute this $P(Y \mid x_1, x_2)$ so we need to apply bay's theorem:

$P(Y \mid x_1, x_2)$ = ($P(Y)$ * $P(x_1, x_2 \mid Y)$)/$P(x_1, x_2)$

so choice number 2 is correct


<!-- grade_id: cell-a84d846ec24f56cd, points: 0, grade: false, solution: true -->
```python
# Provide either True or False as an answer
results = {}
results['3.1.1'] = None
results['3.1.2'] = None
results['3.1.3'] = None

# YOUR CODE HERE
results['3.1.1'] = False
results['3.1.2'] = True
results['3.1.3'] = False
#raise NotImplementedError()
```

<!-- grade_id: cell-b257ff7924dccb48, points: 0.5, grade: true, solution: false -->
```python
assert results.get("3.1.1") is not None
```

<!-- grade_id: cell-3d91bba838237a67, points: 0.5, grade: true, solution: false -->
```python
assert results.get("3.1.2") is not None
```

<!-- grade_id: cell-1830d1023f317f63, points: 0.5, grade: true, solution: false -->
```python
assert results.get("3.1.3") is not None
```

---
## 3.2. Posterior Over Classes with Independence Assumption

Now suppose the features are independent given the class, i.e., $X1 \perp \!\!\! \perp  X2 \mid Y$.
Which of the above three options are sufficient now to compute the posterior distribution $P(Y = 1 \mid x_1, x_2)$? Select all answers that apply!

<div class="alert alert-warning">
Show your calculations in the text field (use markdown & $\LaTeX$) and indicate your solutions in the code cell below. **Keep your solution short**! (2 points)
</div>

<!-- grade_id: cell-7f006e73a5270660, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

since we need to compute this $P(Y \mid x_1, x_2)$ so we need to apply bay's theorem:

$P(Y \mid x_1, x_2)$ = ($P(Y)$ * $P(x_1, x_2 \mid Y)$)/$P(x_1, x_2)$

so choice number 2 is correct

also since $x_1$ and $x_2$ are now independent, we can say that:

$P(x_1, x_2 \mid Y)$ = $P(x_1 \mid Y)$ * $P(x_2 \mid Y)$


so choice 1 can also satisfy bay's theorem so it is correct as well 


<!-- grade_id: cell-61e0215eada549d1, points: 0, grade: false, solution: true -->
```python
# Provide either True or False as an answer
results = {}
results['3.2.1'] = None
results['3.2.2'] = None
results['3.2.3'] = None

# YOUR CODE HERE
results['3.2.1'] = True
results['3.2.2'] = True
results['3.2.3'] = False
#Araise NotImplementedError()
```

<!-- grade_id: cell-b6e374c6c8768319, points: 0.5, grade: true, solution: false -->
```python
assert results.get("3.2.1") is not None
```

<!-- grade_id: cell-dfc3bd41ec7b2773, points: 0.5, grade: true, solution: false -->
```python
assert results.get("3.2.2") is not None
```

<!-- grade_id: cell-92daa65c220d403c, points: 0.5, grade: true, solution: false -->
```python
assert results.get("3.2.3") is not None
```

```python

```

```python

```