```python
import numpy as np
```

# 4. Conditional Independence in Bayesian Networks (6 points)

Consider the following Bayesian network:

<img src="bn.svg" style="height:430px" />

<div class="alert alert-warning">
Indicate which of the following statements are true. $X \perp Y$ means X is independent of Y, $X \perp Y \mid Z$ means X is independent of Y given Z. (1 point / answer).
</div>

Store `True` into the corresponding result variable if the answer must be True and `False` otherwise.

*Hint: Read about **D-separation** and the **Bayes-ball algorithm** if you cannot intuitively solve these problems.* 

 1. $A \perp A$ ?
 2. $M \perp R$ ?
 3. $C \perp I$ ?
 4. $L \perp A \mid D, J$ ?
 5. $H \perp P \mid R$ ?
 6. $K \perp D \mid I, B$ ?

<!-- grade_id: cell-b2674b4fd61a5424, points: 0, grade: false, solution: true -->
```python
 # Provide either True or False as an answer
results = {}
results['4.1'] = None
results['4.2'] = None
results['4.3'] = None
results['4.4'] = None
results['4.5'] = None
results['4.6'] = None

# YOUR CODE HERE
results['4.1'] = False
results['4.2'] = False
results['4.3'] = True
results['4.4'] = True
results['4.5'] = True
results['4.6'] = True
#raise NotImplementedError()
```

<!-- grade_id: cell-c3b644843fab77f4, points: 1, grade: true, solution: false -->
```python
assert results.get("4.1") is not None
```

<!-- grade_id: cell-13f728a3f45423a1, points: 1, grade: true, solution: false -->
```python
assert results.get("4.2") is not None
```

<!-- grade_id: cell-bb6b5105e346f780, points: 1, grade: true, solution: false -->
```python
assert results.get("4.3") is not None
```

<!-- grade_id: cell-e17fd6fbd55cab1c, points: 1, grade: true, solution: false -->
```python
assert results.get("4.4") is not None
```

<!-- grade_id: cell-26465c26053f9953, points: 1, grade: true, solution: false -->
```python
assert results.get("4.5") is not None
```

<!-- grade_id: cell-4e281ffec32cf722, points: 1, grade: true, solution: false -->
```python
assert results.get("4.6") is not None
```

```python

```